---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. dawds

`@instructions`
iajwdiohaefseifjnajlnerbflirnewiufnweufwenhfwuiefbn

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
x = 15
```

`@solution`
```{python}
x = 15
```

`@sct`
```{python}
Ex().check_object("x").has_equal_value()
```
